---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Introduction Data streams

```yaml
type: VideoExercise
key: 7a41f4d29c
xp: 50
```

`@projector_key`
6caad0cb824758b13fb2a1e31d371b08

---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
